# Lab: Geocoding Analysis

```python
import otter
grader = otter.Notebook("geocoding.ipynb")
```

# B. Questions
### **Question 1:** Complete the following code to obtain the coordinates of the Royal Ontario Museum in Canada. 
 - Follow the instructions in the comments.
 - Select only one matching location to the search criteria.

```python
#Remove the ellipsis (...) below and type your code. 

# import the Nominatim geocoder from geopy
...

# Create a Nominatim object in the variable geolocator.  
geolocator = ...

#Geocode the location "Royal Ontario Museum, Canada"
location = ... 

# Assign the values address, latitude, and longitude in the corresponding variables below. 
address = ...
latitude = ...
longitude = ... 

# END SOLUTION
```

```python
grader.check("q1")
```

#### **Question 2:** How can you describe the location geocoded in Question 1? 
- Use the attributes 'type', 'class', and 'address type' to answer this question.
- Respond to the question in a comment 

```python
#Remove the ellipsis (...) below and type your code. 
...
```

```python
grader.check("q2")
```

#### **Question 3:**  Geocode all possible locations that match the criteria "Niagara Falls". 
- For each location, print the name, address, latitude, longitude, and importance.
- Add the addresses in a list  `locations`

```python
#Remove the ellipsis (...) below and type your code. 
...
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

#### **Question 4:** What is the location with the highest importance score? 
Answer in a comment by looking at the results in Question 3. 

```python
#Remove the ellipsis (...) below and type your code. 
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### **Question 5:**  Complete the code below to create a base map pointing at the city of London, Ontario. 
Follow the comments. 

```python
#Remove the ellipsis (...) below and type your code. 
#Remove the ellipsis (...) below and type your code. 
# BEGIN SOLUTION

#Import the folium library. 
...

#Trust this notebook file
!jupyter trust lab11_notebook.ipynb

#Get the coordinates
London_coordinates = ...

#Creating a base map with a focus on London, Ontario
london_map = ...

#Show your map in the cell output. 
...

# END SOLUTION
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6:** Geocode three places around the city of London. You can choose your favourite restaurant, any place at the University that you prefer to study, your favourite shopping center, etc.
- For each place store the name, latitude and longitude.  
- For storing the values, you can use variables. It is not required to use pandas here. It is up to you. 

```python
#Remove the ellipsis (...) below and type your code. 
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7:** Add markers for the three places in the map created in Question 5. Show the map. 

```python
#Remove the ellipsis (...) below and type your code. 
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8:** Save the map created in Question 7. Add any name of your preference. 

```python
#Remove the ellipsis (...) below and type your code. 
...
```

```python
grader.check_all()
```

Tests for q1 (points: 0.1)
```
>>> assert address.startswith('Royal Ontario Museum, 100, Queen')
>>> assert latitude == 43.6676565
>>> assert longitude == -79.3946988
```
success_message: Great! Your code passed this test.


Tests for q2 (points: 0.1)
```
>>> assert location_class == 'tourism'
>>> assert location_type == 'museum'
>>> assert address_type == 'tourism'
```
success_message: Great! Your code passed this test.


Tests for q3 (points: 0.1)
```
>>> assert set(locations) == {'City of Niagara Falls, Niagara County, New York, United States of America', 'Horseshoe Falls, Niagara Falls, Niagara Region, Golden Horseshoe, Ontario, Canada', 'Horseshoe Falls, Niagara Falls, Niagara Region, Golden Horseshoe, Ontario, L2G 7X5, Canada', 'Niagara Falls, City of Niagara Falls, Niagara County, New York, United States of America', 'Niagara Falls, Niagara Region, Golden Horseshoe, Ontario, Canada'}
```
success_message: Great! Your code passed this test.
